---
title: "Filtering Data"
output: html_document
---

## Filtering basics

If we're going to do any kind of data analysis, we better learn how to filter the data that we want to analyze! We'll briefly go over the basics of filtering data in R here. The good news is that most of this maps to your Python/Pandas knowledge.

### Filter with conditions

Suppose we want to filter mtcars for cars with more than 4 **cyl** (cylinders).

<img src="../images/python_logo.png" class="language-icon">

In Python, filtering dataframes can be achieved by placing a boolean expression in a subset expression using `[` 
```{python echo = TRUE, eval = FALSE}
mtcars[mtcars.cyl > 4]
``` 

```{python py-filter}
print(mtcars.cyl[mtcars.cyl > 4])
```

<img src="../images/r_logo.png" class="language-icon">

In R, the same can be achieved using `[` with a condition inside. Luckily, R shares the same comparison operators that you are familiar with in Python.

Try the same filter operation with R:

```{r filter, exercise = TRUE, exercise.lines = 3}

```

```{r filter-hint-1}
"Don't forget that comma for row selection!"
```

```{r filter-hint-2}
mtcars[row_index, ]
```

```{r filter-check}
grade_result(
  pass_if(~ identical(.result, mtcars[mtcars$cyl > 4, ]), "Easy peasy!")
)
```

Besides that pesky comma, this is essentially the gist of filter which is largely similar to how you did it in Python. However, on a later topic, we'll introduce a different way of filtering data and more advanced queries using condition chaining.