---
title: "Filtering Data"
output: html_document
---

## How do I filter data?

### Filter with conditions

Suppose we want to filter mtcars for cars with more than 4 **cyl** (cylinders).

<img src="../images/python_logo.png" class="language-icon">

In Python, filtering dataframes can be achieved by placing a boolean expression in a subset expression using `[` 
```{python echo = TRUE, eval = FALSE}
mtcars[mtcars.cyl > 4]
``` 

```{python py-filter}
print(mtcars.cyl[mtcars.cyl > 4])
```

<img src="../images/r_logo.png" class="language-icon">

In R, the same can be achieved using `[` with a condition inside. Luckily, R shares the same comparison operators that you are familiar with in Python.

Try the same filter operation with R:

```{r filter, exercise = TRUE, exercise.lines = 3}

```

```{r filter-hint-1}
"Don't forget that comma for row selection!"
```
```{r filter-hint-2}
mtcars[row_index, ]
```

```{r filter-check, include = FALSE}
grade_result(
  pass_if(~ identical(.result, mtcars[mtcars$cyl > 4, ]), "Easy peasy!")
)
```