---
title: "group_by"
output: html_document
---

## `group_by`

So far we have been using the tidy verbs to explore the `mtcars` dataframe but we haven't really organized it in terms of a **group** of one or more variables. We'll cover how to do that here. Hold on to your seatbelts! There is a good amount of ground to cover here but it's important to lay down before moving on to the next verb.

### Grouping a dataframe by a variable

Let's group `mtcars` by **cyl** and store the result in a variable called `by_cyl`.

<img src="../images/python_logo.png" class="language-icon">

In Python, we can group a dataframe with a variable using `groupby`:

```{python echo=TRUE, eval=TRUE}
by_cyl = mtcars.groupby('cyl')
print(by_cyl.head())
```
Note how printing the resulting `DataFrameGroupBy` object named `by_cyl` does not *appear* any different from the original dataframe.

<img src="../images/r_logo.png" class="language-icon">

In R, we can use the verb `group_by` to do the same thing. Notice how syntactically similar it is to Python's `groupby`:

```{r group_by-mtcars, echo=TRUE, exercise=TRUE, exercise.lines=3}
by_cyl <- mtcars %>% group_by(cyl)
by_cyl
```
One side-effect of using `group_by` is that the rownames have been silently ignored, but we can always bring that back if we want:

```{r echo=TRUE, context="data"}
by_cyl <- mtcars %>% group_by(cyl) %>% rownames_to_column()
by_cyl
```

Like Python, R doesn't yet reflect any change in the resulting dataframe when printing.

### Operations on a grouped dataframe

The real magic of the resulting object for both Python and R begins when performing operations on it.

#### Slice rows

What if we wanted to slice the resulting grouped dataframe? Let's try to select the first 2 groups based on **cyl**.

<img src="../images/python_logo.png" class="language-icon">

When we want to slice rows of the `GroupBy` object, we have two options in Python.

We can use `head(n)` with some *n* which returns the number of grouped rows we want. For example, if we pass a 1 we would get the first group, i.e. the first set of rows that have 4.0, 6.0 and 8.0 **cyl**. Here's `head(2)`:

```{python echo=TRUE, eval=TRUE}
by_cyl.head(2)
```

Alternately, we could have also used `take` instead and pass in the positional indices.

```{python echo=TRUE, eval=TRUE}
by_cyl.take([0, 1])
```

<img src="../images/r_logo.png" class="language-icon">

Let's see what happens if we use R's `head`:

```{r group_by-slice-head, echo=TRUE, exercise=TRUE, exercise.lines=3}
 head(by_cyl, 2)
```

**Python ⇒ R GOTCHA!** Unlike Python, R's `head` on a grouped dataframe will simply return the number of rows of the dataframe, **NOT** groups of the variable(s)!

If we wanted to mimic Python's `head`, we would instead use the `slice` verb. Here we grab 2 groups:

```{r group_by-slice-slice, echo=TRUE, exercise=TRUE, exercise.lines=3}
slice(by_cyl, 1:2)
```

From this, it would seem that `slice` bears more resemblance to Python's `head` than `take` except that `slice` requires a range.

However, watch what happens if you we had just supplied 2 to `slice`:

```{r group_by-slice-single, echo=TRUE, exercise=TRUE, exercise.lines=3}
slice(by_cyl, 2)
```
**Python ⇒ R: GOTCHA!**  Unlike Python's `head`, supplying a number for R's `slice` on a grouped dataframe will return the group corresponding to the positional index supplied, **NOT** groups of the variable(s)!

The `slice` now behaves like Python's `take`:

```{r group_by-slice-single-2, echo=TRUE, exercise=TRUE, exercise.lines=3}
slice(by_cyl, c(1, 2))
```

Look familiar?

<img src="../images/python_logo.png" class="language-icon">
```{python echo=TRUE, eval=TRUE}
by_cyl.take([0, 1])
```

So, if you want to mimic Python's `head` on a grouped dataframe, use R's `slice` with a range. But, if you want to mimic Python's `take`, you can use the `slice` verb and supply either a single positional index or a list of indices using `c()`.

To stay consistent, we'll just use `take` for the Python and `slice` for R from here on out.

#### Select columns

What if we wanted to select the **mpg** and the **disp** columns on the grouped dataframe?

<img src="../images/python_logo.png" class="language-icon">

To select rows of the `GroupBy` in Python, we can use the `[[`. But, we would have to use take so that **cyl** gets included in the result. We'll print the first two groups like before:

```{python echo=TRUE, eval=TRUE}
by_cyl[['mpg', 'disp']].take([0,1])
```
Python doesn't retain the **cyl** group variable when using `head` even though it's there internally so we use `take`.

<img src="../images/r_logo.png" class="language-icon">

We can use R's `select` verb on grouped data to select columns. No big surprises here:

```{r group_by-select-simple, echo=TRUE, exercise=TRUE, exercise.lines=5}
by_cyl %>%
  select(mpg, disp) %>%
  slice(1:2)
```
Unlike Python, R notices that you are operating on a grouped dataframe, so it includes the **cyl** even if not explicitly selecting it. But, we can also explicitly select with a convenience function in `dplyr` called `group_cols`:

```{r group_by-select-grouped, echo=TRUE, exercise=TRUE, exercise.lines=5}
by_cyl %>%
  select(group_cols(), mpg, disp) %>%
  slice(1:2)
```

If we want to keep the rownames, it takes a bit more typing but it can be done. You need to start a fresh tidy chain from `mtcars` and repeat the `group_by` unfortunately:

```{r group_by-select-group, echo=TRUE, exercise=TRUE, exercise.lines=6}
mtcars %>%
  group_by(cyl) %>%
  select(group_cols(), mpg) %>%
  rownames_to_column()
```

However, slicing and selecting columns on grouped dataframes isn't all that interesting. Let's look at what else we can do.

#### Filter rows

More interestingly, we can also filter the grouped data. Let's filter for rows where we get the max **disp** (displacement or volume of all the cylinders) for each **cyl**.

<img src="../images/python_logo.png" class="language-icon">

We can use the `idxmax` in Python to first find the max **disp** for each group, then use `loc` to index those particular rows from `mtcars`:

```{python echo=TRUE, eval=TRUE}
mtcars.loc[by_cyl.disp.idxmax()]
```

<img src="../images/r_logo.png" class="language-icon">

In R, we can simply pipe the grouped dataframe to `filter` and filter for those rows where **disp** is the max.

```{r group_by-filter, echo=TRUE, exercise=TRUE, exercise.lines=3}
# for each cyl, give me row of the max disp
by_cyl %>% filter(disp == max(disp))
```
Unlike Python, we didn't have to bring back the `mtcars` to perform the filtering. The other small difference in output here is that R does not sort by **cyl** automatically like Python. But you could do so using `arrange`, which brings us to our next verb operating on a grouped dataframe.

#### Arrange

Let's sort `mtcars` according to **disp** for each **cyl**.

<img src="../images/python_logo.png" class="language-icon">

We can use the `apply` in Python to sort by supplying a lambda function:

```{python echo=TRUE, eval=TRUE}
by_cyl.apply(lambda x: x.sort_values(by='disp', ascending=False))
```

<img src="../images/r_logo.png" class="language-icon">

For R, things are much easier! We simply pipe the grouped data to our `arrange` verb and use `desc` to sort:

```{r group_by-arrange-nogroup, echo=TRUE, exercise=TRUE, exercise.lines=3}
by_cyl %>% arrange(desc(disp))
```

Wait a minute. It looks like **cyl** also got sorted by descending order! We'd like to keep **cyl** stay in the natural 4, 6, 8 cylinders group order.

**Base R ⇒ Tidy R: GOTCHA!** Grouped arrange ignores groups by default. Note how the **cyl** also gets sorted when we just want to sort by **disp**.

To fix this, you specifically ask to sort by group by setting the `.by_group` to `TRUE`:

```{r group_by-arrange-group, echo=TRUE, exercise=TRUE, exercise.lines=3}
by_cyl %>% arrange(desc(wt), .by_group = TRUE)
```

#### Mutate

The very last `group_by` + verb interaction we will look at is for `mutate`. To demonstrate this interaction, we'll do a rather silly example, but should be good enough to show the difference between Python and R. 

Let's create a new column `rank` that represents the rank of **mpg** for each group of **cyl**, where a rank of 1 would represent the highest mpg. It will be easy to check if we did this right since the order should be the same as if we were sorting **mpg** by descending order!

<img src="../images/python_logo.png" class="language-icon">

In Python we actually can't add a new column to a `GroupBy`. But, we can first create the **rank** column which is the ranking of **mpg** where the highest value is rank 1 and so on. Then, we can use the groupby with `apply` where the lambda function sorts according to **rank** in descending order:

```{python echo=TRUE, eval=TRUE}
mtcars['rank'] = mtcars.mpg.sort_values().rank(method='min', ascending=False)
mtcars.groupby('cyl').apply(lambda x: x.sort_values(by='rank'))
```

<img src="../images/r_logo.png" class="language-icon">

Unlike Python, R *does* allow us to add a new column to the grouped dataframe using `mutate` allowing us to pipe the grouped dataframe right to `mutate` to first create the **rank** column. Then we can use another handy function from `dplyr` called `min_rank` which ranks the **mpg** column:

```{r group_by-mutate-simple-arranged, echo=TRUE, exercise=TRUE, exercise.lines=5}
by_cyl %>% 
  mutate(rank = min_rank(desc(mpg)))
```

<!-- You need to be careful when using `mutate` to overwrite columns. Let's try doing that for our grouping variable **cyl** by executing the following. -->

<!-- ```{r group_by-mutate-bad, echo=TRUE, exercise=TRUE, exercise.lines=3} -->
<!-- by_cyl %>% mutate(gdf, cyl = cyl * 2) -->
<!-- ``` -->

<!-- **GOTCHA!** When you group variables with `group_by`, you can't overwrite the grouping variables using `mutate`! You should've seen this error: -->

<!-- ``` -->
<!-- Error : Column `cyl` can't be modified because it's a grouping variable -->
<!-- ``` -->

Almost there. We need to now order it by **cyl** to mimic the Python output.

Let's finish off the code by doing the following:

- Move the **rank** column up to the front so we can easily see it
- Finally `arrange` the rows by rank.

**Note**: You don't need to preserve row names and you can print the whole resulting dataframe.

```{r group_by-mutate-chained-exercise, echo=TRUE, exercise=TRUE, exercise.lines=6}
by_cyl %>% 
  mutate(rank = min_rank(desc(mpg)))



```

```{r group_by-mutate-chained-exercise-hint-1}
"Use a chain of piped calls using %>%"
```

```{r group_by-mutate-chained-exercise-hint-2}
"Don't forget about that gotcha on `arrange`!"
```
```{r group_by-mutate-chained-exercise-hint-3}
by_cyl %>% 
  mutate(rank = min_rank(desc(mpg))) %>% 
  select(...) %>% 
  arrange(...)
```

```{r group_by-mutate-chained-exercise-check}
grade_result(
  pass_if(~ identical(
    .,
    by_cyl %>% 
      mutate(rank = min_rank(desc(mpg))) %>% 
      select(rank, everything()) %>% 
      arrange(rank, .by_group = TRUE)
    )
  )
)
```

Once you've done that we can clearly see the rows ordered according to the rank which is in turn based on **mpg** for each **cyl** groups 4, 6, and 8.

For most data operations, we operate on groups of variables, so `group_by` is an important addition to our arsenal! This will now set us up for performing basic statistics on the data on the next topic with a new verb called `summarise`.

